---
title: "Tutorial"
output: learnr::tutorial
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(ggplot2)
library(palmerpenguins)
library(DiagrammeR)
knitr::opts_chunk$set(echo = FALSE)

```

## `grade_this_code()`

Estamos probando la funci贸n `grade_this_code()`

```{r ejercicio1, exercise = TRUE, exercise.cap = "Usamos correct/incorrect"}
#Suma 2 + 24

```

```{r ejercicio1-solution}
2+24
```

```{r ejercicio1-check}
gradethis::grade_this_code()
```

Este ejemplo nos hace ver que no es demasiado estricto con los espacios o el lugar donde se ponga la respuesta (dentro del ejercicio chunk)

```{r ejercicio2, exercise = TRUE, exercise.cap = "Usamos allow_partial_matching"}
#Haz una muestra de 10 normales con media 3 y varianza 2

```

```{r ejercicio2-solution, exercise.reveal_solution = FALSE}
rnorm(10, mean = 3, sd = 2)
```

```{r ejercicio2-check}
gradethis::grade_this_code(correct = "Sabes leer", incorrect = "No sabes nada", allow_partial_matching = TRUE)
```

Nos aparece un mensaje de "error" cuando usamos `allow_partial_matching = TRUE` (si es que no escribimos bien alg煤n caracter de la solution) pero no podemos hasta ahora modificar tal mensaje.

Aparece correcta la opci贸n aun cuando dentro de la funci贸n (`rnorm`) no especificamos el nombre del argumento sino solo el valor que deseamos que tome.

```{r ejercicio3, exercise = TRUE, exercise.cap = "Usando mock", exercise.lines = 8, message=FALSE, warning=FALSE}
# Completa el sieguiente c贸digo con los datos de c贸digo:
grafica_pinguinos <- ggplot(data = ____) +
 aes(x = flipper_length_mm) +
 geom_histogram(aes(fill = species),
                     alpha = 0.5,
                     position = "identity") 

grafica_pinguinos
```

```{r ejercicio3-solution, exercise.reveal_solution = FALSE}
# Completa el siguiente c贸digo:
grafica_pinguinos <- ggplot(data = penguins) +
 aes(x = flipper_length_mm) +
 geom_histogram(aes(fill = species),
                     alpha = 0.5,
                     position = "identity") 

grafica_pinguinos
```

```{r ejercicio3-check}
gradethis::grade_this_code(correct = "Felicidades", incorrect = "No felicidades")
```

## `grade_this()`

La funci贸n `grade_this()` remplaz贸 a la funci贸n `grade_result()`. Ambas evaluan solo el resultado del ejercicio con el resultado de la soluci贸n sin comparar el c贸digo. Sin embargo, `grade_this()` puede ser personalizada para definir lo que queremos evaluar en cada ejercicio.

```{r ejercicio4, exercise = TRUE, exercise.cap = "Con grade_this()"}
#Suma 2 + 24
```

```{r ejercicio4-solution}
2+24
```

```{r ejercicio4-check}
#Obtener la media de la longitud de las alas de los pinguinos del dataset penguins
gradethis::grade_this(
  if(.result == .solution){
    pass("Muy bien")
  } else{
    fail("Intenta usar na.rm   = TRUE para el c谩lculo")
  }
)
```

Al usuar `.result` hace referencia al resultado que se calcule en el chunk donde el usuario escribi贸.

Por otro lado `.solution` considera lo que resulta del chunk-solution.

```{r diagrama}
DiagrammeR("
A[gradethis] --> B[grade_this_code()];
A --> C[grade_this()];
B --> D[Correct];
B --> E[Incorrect];
B --> F[allow_partial_matching];
C --> G[pass(), fail(), .result, .solution];

")
```

# RESTRINGIR LA DE EJERCICIOS 

## Calificar ejercicios de c贸digo 

Si deseas que los usuarios tengan retroalimentaci贸n cuando cometan alg煤n error en los ejercicios de c贸digo (en t茅rminos del objetivo de aprendizaje en el ejercicio), puedes implementarlo con lo que veremos en esta secci贸n.

### Instalar `gradethis`

La paqueter铆a que usaremos para este fin se llama `gradethis` y te sugerimos hacer lo siguiente para tenerla instalada en tu computadora:

1.  *Si ya instalaste `learnr` con el comando `install.packages("learnr")`* entonces utiliza `remove.packages("learnr")` para desinstalarlo.

    Sea afirmativo o no lo anterior, haz lo siguiente:

2.  Usa el comando `remotes::install_github("rstudio/learnr")` para tener la versi贸n en desarrollo de `learnr`.

3.  Usa el comando `remotes::install_github("rstudio-education/gradethis", force = TRUE)` para instalar `gradethis`.

4.  Cargar la paqueter铆a con el comando `library(gradethis)` despu茅s de cargar `learnr`, ambas dentro del chunk inicial (que se puede llamar setup).

### Esquema

![](images/gradethis.png){width="673"}

Existen dos formas principales para evaluar el ejercicio de c贸digo:

1.  Evaluar el c贸digo.

    En la paqueter铆a `gradethis` esto se puede hacer con la funci贸n `grade_this_code()`. Esta funci贸n reemplaza a la funci贸n `grade_code()` ( que se menciona [en la documentaci贸n de `learnr`](https://rstudio.github.io/learnr/exercises.html#Exercise_Checking)).

2.  Evaluar el resultado que se deriva del c贸digo.

    Con la funci贸n `grade_this()` se puede personalizar el criterio de evaluaci贸n con respecto al resultado. **Por ejemplo**: si el resultado es mayor al esperado, se puede implementar unmensaje que diga "el resultado correcto es menor al obtenido" .

    Esta funci贸n, como en el punto anterior, reemplaza a `grade_result()`.

### grade_this_code()

**Ejemplo:**

Implementemos un ejercicio donde el usuario tenga que sumar 2+2 (= ) en el que se evaluar谩 que el c贸digo sea `2+2`.

1.  Crear el chunk de ejercicio.

2.  Crear el chunk solution.

    ```{r gejercicio-solution, exercise.reveal_solution = FALSE}
    2+2
    ```

3.  Crear el chunk check:

    El chunk de ejercicio se llama `gejericicio`, por lo que el nombre del chunk check debe ser `gejercicio-check`.

    Dentro del chunk siguiente se usar谩 la funci贸n `grade_this_code()`.

        '''{r gejercicio-check}

        grade_this_code()

        '''

    ```{r gejercicio-check}
    gradethis::grade_this_code()
    ```

Resultado:

```{r gejercicio, exercise = TRUE, exercise.cap = "Suma"}
# Suma dos m谩s dos

```

-   Nota que no importa si pones espacios en el c贸digo o pones el c贸digo en otras l铆neas (diferentes a las del chunk solution).

-   Nota que se a帽ade un bot贸n que se llama `Submit answer` para que el c贸digo sea evaluado.

-   Intenta escribir `3+1` , `2*2`, `4*1,` ... y nota que esa no es la respuesta correcta aun cuando el resultado sea el mismo (4).

### Par谩metros

Se pueden modificar los mensajes que aparecen cuando se comete un error o un acierto con los par谩metros de `grade_this_code()`:

-   `correct` recibe una cadena de caracteres la cual aparece cuando el c贸digo del usuario y el c贸digo de la soluci贸n coinciden.

    Por ejemplo `correct = "Haz hecho la suma que deseo"`

-   `incorrect` recibe una cadena de caracteres la cual aparece cuando el c贸digo del usuario y el c贸digo de la soluci贸n NO coinciden.

    Por ejemplo `incorrect = "Intenta sumar dos n煤meros iguales y que el resultado sea 4"`

-   `allow_partial_matching` puede ser `TRUE` o `FALSE`. Por defecto es `FALSE` y cuando `allow_partial_matching = TRUE` entonces se permite que aparezcan mensajes que indican d贸nde no coincide el c贸digo con la soluci贸n.

**Ejemplo** del ejemplo anterior tenemos:

    '''{r gejercicio-check}

    grade_this_code(correct = "Muy bien", 
                    incorrect = "Intenta sumar dos n煤meros iguales y que el resultado sea 4", 
                    allow_partial_matching = TRUE)

    '''

```{r gejercicio1, exercise = TRUE, exercise.cap = "Suma"}
# Suma dos m谩s dos

```

```{r gejercicio1-solution, exercise.revela_solution=FALSE}
2+2
```

```{r gejercicio1-check}
grade_this_code(correct = "Muy bien", 
                incorrect = "Intenta sumar dos n煤meros iguales y que el resultado sea 4", 
                allow_partial_matching = TRUE)
```

### 驴Te qued贸 duda?

```{r gcuestionario}
quiz(
  question("Si el chunk del ejercicio se llama 'ejercicio', 驴En qu茅 chunk debe ir la funci贸n `grade_this_code()`?",
    allow_retry = TRUE,
    answer("ejercicio-solution", message = "No, aqu铆 debe ir el c贸digo esperado"),
    answer("ejercicio-check", correct = TRUE),
    answer("ejercicio-cechk", message = "Revisa la escritura"),
    correct = "隆S铆 que has puesto atenci贸n!",
    try_again = "Vuelve a leer la secci贸n anterior",
    submit_button = "Yo s茅 que s茅",
    try_again_button = "Int茅ntalo de nuevo "
  ),
  
  question("驴Qu茅 par谩metro permite al usuario saber en qu茅 parte del c贸digo se pudo haber equivocado?", 
    answer("allow_retry"),
    answer("allow_partial_matching", correct = TRUE),
    answer("incorrect"),
    allow_retry = TRUE,
    correct = "隆S铆 que has puesto atenci贸n!",
    try_again = "Vuelve a leer la secci贸n anterior",
    submit_button = "Yo s茅 que s茅",
    try_again_button = "Int茅ntalo de nuevo "
  ),
  
  question("驴Para que sirve el par谩metro correct?",
    answer("Para permitir que el usuario tenga correcto lo que haga siempre"),
    answer("Para dar a conocer al usuario un mensjae cuando el su c贸digo coincide con el de la soluci贸n", correct = TRUE),
    answer("Para estar felices"),
    allow_retry = TRUE,
    correct = "隆S铆 que has puesto atenci贸n!",
    try_again = "Vuelve a leer la secci贸n anterior",
    submit_button = "Yo s茅 que s茅",
    try_again_button = "Int茅ntalo de nuevo "
  ), 
  caption = "Pongamos a prueba lo que haz aprendido"
)
```
