---
title: "Tutorial de learnr"
author: 
  - name: "Yanely L.G."
  - name: "Carlos A. Ar."
  - name: "Facultad de Ciencias"
  - name: "UNAM"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: "Este es un tutorial de learnr para aprender a usar learnr."
---

```{r setup, include=FALSE}
#Librer铆as a usar
library(learnr)
library(gradethis)
library(ggplot2)
library(palmerpenguins)
library(DiagrammeR)
#Modificaci贸nes para los chunks
knitr::opts_chunk$set(echo = FALSE)
#Modificaciones en el tutorial
```

```{=html}
<style>
div {
  text-align: justify;
}
/***********************/
/*T铆tulos ## y autor铆a F6EACD*/
/***********************/
h1, h2 {
  color: #6D4101;
  padding: .4em; /*es decir el 60 por ciento del tama帽o de la letra*/
  font-family: "Courier New", monospace;
  font-size: 35px;
  font-weight: bold;
  opacity: 0.95;
}
em {
  color: #8A6C0C;
  font-family: "Courier New", monospace;
  font-size: 18px;
  padding: .2em;
  opacity: 0.97;
  font-weight: bold;
}
/***********************/
/* Cuerpo */
/***********************/
body {
  font-family: Verdana, Helvetica, sans-serif;
  font-size: 16;
  background-color: #EEF5F9;
  color: #000000;
}
/***********************/
/* T铆tulos con ###     */
/***********************/
h3, h4, h5{
  color: #06116C;
  font-family: Verdana, Helvetica, sans-serif;
}
/*******************************************/
/* Referencias a p谩ginas o a otro documento*/
/*******************************************/
a {
    color: #B8870F;
    text-decoration: none;
}
/**********************/
/* Chunks de exercises*/
/**********************/
.ace-tm {
    background-color: #FFFEF7;
    color: #010730;
}
/***********************/
/* tabla de contenidos */
/***********************/
.topicsList {
  padding: .5em;
}
.topicsHeader {
  color: #06116C;
  padding: .5em;
}
.topicsList #doc-metadata {
  color: #06116C;
  padding: .5em;
}
.topicsList .topic.current {
  background-color: #8A6C0C;
  color: #FFFFFF;
  font-weight: bold;
}
.topicsList .topic:hover, .topicsList .topic:active {
  background-color: #E0CD4E;
  color: #01053D;
}
/***********/
/* Botones */
/***********/
.btn {
  background-color: #000857;
  color: #EBEDFF;
}
/* cambiar de p谩gina */
.btn-default {
    color: #EBEDFF;
    background-color: #000857;
    border: none;
}
/* cambiar de p谩gina */
.btn-light {
  background-color: #000857;
  color: #EBEDFF;
}
/* run code */
.btn-primary , .btn-success, .btn-info{
  background-color: #786300;
  color: #EBEDFF;
}
/* mientras est谩 desabilitados */
.btn:hover, .btn:active, .btn:disabled {
  background-color: #786300;
  color: #EBEDFF;
}
/**********/
/* C贸digo */
/**********/
code {
    color: #6D4101;
    background-color:  #F7FBF9;
    font-size: 15px;
    font-weight: bold;
}
/**************************/
/* Recuadro para ejemplos */
/**************************/
.boxed {
    background: #F7FBF9;
    color: black;
    border: 3px solid #C83737;
    margin: 0px auto;
    width: 456px;
    padding: 10px;
    border-radius: 10px;
  }
.note {
    padding: 1em;
    margin: 1em 0;
    padding-left: 100px;
    background-size: 70px;
    background-repeat: no-repeat;
    background-position: 15px 15px;
    min-height: 120px;
    color: black;
    background-color: lightgrey;
    border: solid 5px #C83737;
    background-image: url("manzana.png");
  }
.note_white {
    padding: 1em;
    margin: 1em 0;
    padding-left: 100px;
    background-size: 70px;
    background-repeat: no-repeat;
    background-position: 15px 15px;
    min-height: 120px;
    color: black;
    background-color: #F7FBF9;
    border: solid 5px #C83737;
    background-image: url("manzana.png");
  }
</style>
```
## Introducci贸n

### Objetivo de este tutorial

El objetivo de este tutorial es que aprendas a generar tu propio tutorial () interactivo en R usando la paqueter铆a `learnr`. Esta paqueter铆a consta de funciones que permiten modificar un [archivo .Rmd](https://tutorialesmg.netlify.app/r-markdown) para incluir cuestionarios, ejercicios de c贸digo entre otras cosas.

### Esquema de este tutorial

```{r, echo=FALSE, eval = FALSE}
DiagrammeR("
graph LR;
  A[learnr] --> B[Im谩genes y Videos];
  A --> D[Cuestionarios];
  A --> C[Ejercicios de c贸digo];
  A --> E[Dise帽o del tutorial];
  C --> F[Opciones de C贸digo];
  C --> H[solution y hints];
  C --> Q[Calificar ejercicios de c贸digo];
  F --> I[exercise.cap];
  F --> J[exercise.completion];
  F --> K[exercise.setup];
  D --> L[quiz];
  D --> M[question];
  D --> N[answer];
  E --> O[CSS];
  E --> P[theme];
  Q --> R[gradethis];
  
  style A fill:#EAE38F, stroke:#A59800;  
  
  style B fill:#FCF3CF, stroke:#F0B27A; 
  style C fill:#73C6B6, stroke:#138D75;
  style D fill:#C39BD3, stroke:#A569BD;  
  style E fill:#AED6F1, stroke:#5499C7;
  
  style F fill:#A9DFBF, stroke:#82E0AA;  
  style H fill:#A9DFBF, stroke:#82E0AA; 
  style Q fill:#A9DFBF, stroke:#82E0AA; 
  
  style L fill:#E8DAEF, stroke:#D2B4DE;  
  style M fill:#E8DAEF, stroke:#D2B4DE;
  style N fill:#E8DAEF, stroke:#D2B4DE;  
  
  style O fill:#EBF5FB, stroke:#7FB3D5;
  style P fill:#EBF5FB, stroke:#7FB3D5;
  
  style I fill:#D0ECE7, stroke:#73C6B6;  
  style J fill:#D0ECE7, stroke:#73C6B6;
  style K fill:#D0ECE7, stroke:#73C6B6;
  style R fill:#D0ECE7, stroke:#73C6B6;")
```

![](images/esquema.png)

### Primeros pasos

![](https://www.youtube.com/watch?v=EBAvQkP9Kyc)

```{r cuestionario1}
quiz(
  question(
    text = "驴Qu茅 paqueter铆a se requiere para crear un tutorial interactivo en RStudio?",
    allow_retry=TRUE,
    answer("xaringan",message = "Esta paqueter铆a sirve para crear presentaciones en R."),
    answer("flipbookr",message = "Esta paqueter铆a sirve para mostrar c贸digo incremental."),
    answer("learnr", correct = TRUE, message = "隆Correcto!"),
    correct = "隆Felicidades, vas al camino del 茅xito!",
    try_again = "T煤 puedes mejorar",
    submit_button = "Confirmar",
    try_again_button = "Intenta de nuevo"
  ),
  question(
    text = "驴Qu茅 necesitas para crear un tutorial interactivo?",
    allow_retry=TRUE,
    answer("Determinar los objetivos del tutorial",correct=TRUE),
    answer("La paqueter铆a learnr",correct=TRUE),
    answer("Saber matem谩ticas"),
    correct = "隆Felicidades!",
    try_again = "T煤 puedes mejorar",
    submit_button = "Confirmar",
    try_again_button = "Intenta de nuevo"
  ),
  caption = "Cuestionario"
)

```

## Crear un tutorial interactivo

### Instalar y cargar `learnr`

Usa este comando `install.packages("learnr")` en tu consola de R/RStudio para instalar el paquete `learnr`.

Como ejercicio, ingresa el comando para cargar la paqueter铆a dentro del archivo de tu tutorial.

```{r ej1, exercise=TRUE,exercise.cap="Cargar learnr"}

```

```{r ej1-solution}
library(learnr)
```

```{r ej1-check}
gradethis::grade_this_code(correct = "Excelente", incorrect = "Usa la funci贸n library")
```

### Crear un tutorial

**Opci贸n 1: Para crear el tutorial desde cero**

1.  Crear un nuevo Rmarkdown:

    ![](images/learnr1.png){width="489"}

2.  Seleccionar la plantilla de **Interactive tutorial**

    ![](images/learnr2.png){width="473"}

**Opci贸n 2: Modificar un .Rmd ya existente**

1.  En un archivo Rmarkdown modificar el [YAML](https://tutorialesmg.netlify.app/r-markdown) header para incluir el output que se muestra:

<!-- -->

    ---
    title: "Tutorial"
    output: 
      learnr::tutorial:
        theme: flatly
        highlight: tango
    runtime: shiny_prerendered
    description: "Esto es un tutorial de learnr para aprender a usar learnr."
    ---

## Ejercicios de c贸digo {data-allow-skip="FALSE"}

**驴Te gustar铆a que los estudiantes puedan aplicar lo que se ense帽a en tu tutorial sin la necesidad de abrir otro archivo o la consola de R?** En esta secci贸n aprenderemos a crear ejercicios de c贸digo, los cuales permiten ejecutar c贸digo de R dentro del tutorial.

### 驴C贸mo crear un ejercicio de c贸digo?

1.  Crear un chunk de c贸digo R. En RStudio podemos usar el atajo `CTRL + ALT + I`.

2.  A帽adir la opci贸n `exercise=TRUE` en el chunk para convertirlo en un ejercicio.

Ejemplo:

C贸digo

```{r eval=FALSE,echo=TRUE}
'''{r ejercicio, exercise=TRUE}
#Aqu铆 puedes ingresar c贸digo de R y ejecutarlo.
'''
```

Resultado

```{r ejercicio0, exercise=TRUE}
#Aqu铆 puedes ingresar c贸digo de R y ejecutarlo.
```

### 驴C贸mo modificar un ejercicio de c贸digo?

Para modificar un ejercicio de c贸digo usamos las opciones del chunk del ejercicio, las cuales no se muestran en el tutorial (y en general no se muestran en los .Rmd).

    ```{r nombre_del_chunk, diversas_opciones_que_mostraremos}
    # Tu c贸digo
    ```

Ejemplo:

![](images/learnr16.png)

Veamos qu茅 es posible hacer con los ejercicios de c贸digo usando las opciones de la siguiente tabla:

+---------------------------------------------------------------------+------------------------+--------------------------------------------------------------------------------------------+--------------------------------------------------------------+
| 驴Quieres...                                                         | Usa la opci贸n:         | 驴Qu茅 valores toma la opci贸n?                                                               | Resultado                                                    |
+=====================================================================+:=======================+:===========================================================================================+==============================================================+
| cambiar el n煤mero de lineas que aparecen en el ejercicio?           | `exercise.lines = n`   | n puede ser: 1, 2, 3, .... Es el n煤mero de filas de c贸digo que se muestren en el tutorial. | ![](images/learnr5.png){width="281"}                         |
+---------------------------------------------------------------------+------------------------+--------------------------------------------------------------------------------------------+--------------------------------------------------------------+
| modificar el t铆tulo del ejercicio?                                  | `exercise.cap`         | Una cadena de caracteres. Por default, `exercise.cap = "Code"`                             | ![](images/learnr6.png){width="281"}                         |
+---------------------------------------------------------------------+------------------------+--------------------------------------------------------------------------------------------+--------------------------------------------------------------+
| preevaluar el c贸digo del ejercicio?                                 | `exercise.eval`        | Un boleano. Por default es `FALSE`.                                                        | ![](images/learnr7.png){width="281"}                         |
+---------------------------------------------------------------------+------------------------+--------------------------------------------------------------------------------------------+--------------------------------------------------------------+
| evitar que el c贸digo ingresado tome demasiado tiempo en ejecutarse? | `exercise.timelimit`   | Un entero. Representa el tiempo l铆mite de ejecuci贸n del c贸digo en segundos.                | ![](images/learnr8.png){width="281"}                         |
|                                                                     |                        |                                                                                            |                                                              |
|                                                                     |                        | **No significa** que el usuario tenga ese tiempo para completar el c贸digo.                 |                                                              |
+---------------------------------------------------------------------+------------------------+--------------------------------------------------------------------------------------------+--------------------------------------------------------------+
| brindar o quitar ayuda al estudiante?                               | `exercise.completion`  | Un boleano. Por default es `TRUE`.                                                         | ![](images/learnr9.png){width="281"}                         |
+---------------------------------------------------------------------+------------------------+--------------------------------------------------------------------------------------------+--------------------------------------------------------------+
| prevenir al estudiante sobre errores de codificaci贸n?               | `exercise.diagnostics` | Un boleano. Por default es `TRUE`.                                                         | Ve ejemplo en la secci贸n: **驴C贸mo puedo ayudar al usuario?** |
+---------------------------------------------------------------------+------------------------+--------------------------------------------------------------------------------------------+--------------------------------------------------------------+
| permitir que se reinicie el ejercicio?                              | `exercise.startover`   | Un boleano. Por default es `TRUE`.                                                         | ![](images/learnr10.png){width="281"}                        |
+---------------------------------------------------------------------+------------------------+--------------------------------------------------------------------------------------------+--------------------------------------------------------------+

El uso de estas opciones depende del prop贸sito del tutorial, pero es posible que se requiera la misma opci贸n para todos los ejercicios. **驴Existe alguna forma de fijar opciones para todos los chunks de ejercicio a la vez?**

**隆S铆!** con la funci贸n `tutorial_options()` podemos ahorrar tiempo al fijar dichas opciones. Esta configuraci贸n global puede especificarse en un chunk con nombre `setup` al inicio del tutorial. Tambi茅n es recomendable cargar todas las librer铆as que se requieran para el tutorial en dicho chunk.

**驴C贸mo funciona un chunk de ejercicio?** Cada chunk de ejercicio tiene su propio ambiente en R, por lo que los objetos creados en un ejercicio no estar谩n disponibles para cualquier otro chunk.

Una forma permitir que dichos objetos est茅n disponibles fuera del ejercicio en el que se crean es usando la opci贸n `exercise.setup="nombre-chunk-a-usar"`en el chunk de ejercicio en el que queremos usar el resultado previo. Podemos ir formando una cadena de resultados a usar de esta forma.

**Ejemplo.**

    {r respuesta}
    grafica_pinguinos <- ggplot(data = penguins) +
           aes(x = flipper_length_mm) +
      geom_histogram(aes(fill = species),
                     alpha = 0.5,
                     position = "identity") 
                     
                    
    1.

    {r ejercicio7, exercise=TRUE, message=FALSE,warning=FALSE}
    grafica_pinguinos <- ggplot(data = ______) +
           aes(x = flipper_length_mm) +
      geom_histogram(aes(fill = species),
                     alpha = 0.5,
                     position = "identity") +
      labs(x = "Flipper length (mm)",
           y = "Frequency",
           title = "Penguin flipper lengths") + 
      scale_fill_manual(values = c("darkorange","purple","cyan4"))
      
      grafica_pinguinos
      
    2. 

    {r ejercicio7_2, exercise=TRUE, exercise.setup="respuesta"}
    grafica_pinguinos + _____

1.  Completa el siguiente c贸digo para generar una gr谩fica con los datos del data.frame `penguins`

```{r respuesta}
grafica_pinguinos <- ggplot(data = penguins) +
  aes(x = flipper_length_mm) + 
  geom_histogram(aes(fill = species), alpha = 0.5, position = "identity") 
```

```{r ejercicio7, exercise=TRUE, message=FALSE,warning=FALSE}
grafica_pinguinos <- ggplot(data = ____) +
    aes(x = flipper_length_mm) +
    geom_histogram(aes(fill = species),
                     alpha = 0.5,
                     position = "identity") 
grafica_pinguinos
```

<!-- -->

2.  Agrega el tema `minimal` al gr谩fico anterior.

```{r ejercicio7_2, exercise=TRUE, exercise.setup="respuesta",message=FALSE,warning=FALSE}
grafica_pinguinos + _____

```

```{r ejercicio7_2-solution}
grafica_pinguinos  + theme_minimal()

```

*Notemos que la segunda parte del ejercicio va ligada a un chunk que contiene la respuesta correcta y no al chunk de la parte 1 del ejercicio (que contiene la respuesta del usuario).* 

### Ejercicios con hints y soluciones {#syh}

Si alg煤n ejercicio de tu tutorial requiere una muy buena comprensi贸n del c贸digo o consideras que tiene un nivel alto de dificultad, **puedes incluir hints para los estudiantes (sujerencias/orientaci贸n sobre c贸mo resolver el ejercicio).** Adem谩s puedes **incluir el c贸digo de la soluci贸n para que los estudiantes puedan identificar la forma correcta de resolver el ejercicio.**

Para lograr esto usamos el sufijo `-hint` en el nombre del chunk que contiene el hint o mostrar el c贸digo de la soluci贸n a帽adiendo `-solution` al nombre del chunk que contiene la soluci贸n.

Por ejemplo, para proporcionar c贸digo que de una pista para resolver el ejercicio que tiene nombre de chunk `ejercicio1`, escribimos la pista en un chunk que tenga nombre `ejercicio1-hint`. De manera similar, para proporcionar el c贸digo de la soluci贸n de este ejercicio lo escribimos en un chunk con nombre `ejercicio1-solution`.

C贸digo:

    Este es el c贸digo del chunk del ejercicio:

      '''{r ejercicio1, exercise=TRUE}
      #Aqu铆 va el ejercicio:
      
      # Cambia al color de los puntos a rosa
      ggplot(mtcars, aes(x=mpg,y=cyl)) +
        geom_point()
      '''
    Este es el c贸digo del chunk del hint:

      '''{r ejercicio1-hint}
      #Aqu铆 va el hint para el ejercicio 1
      
      # Usa el par谩metro color dentro de geom_point()
      '''

Resultado:

```{r ejercicio1, exercise=TRUE}
# Cambia al color de los puntos a rosa
ggplot(mtcars, aes(x=mpg,y=cyl)) +
  geom_point()
```

```{r ejercicio1-hint}
#Usa el par谩metro color dentro de geom_point()
```

C贸digo:

    Este es el c贸digo del chunk del ejercicio:

      '''{r ejercicio2, exercise=TRUE}
      #Aqu铆 va el ejercicio:
      
      # Cambia al color de los puntos a verde
      ggplot(mtcars, aes(x=mpg,y=cyl)) +
        geom_point()
      '''
    Este es el c贸digo del chunk de la soluci贸n:

      '''{r ejercicio2-solution}
      #Aqu铆 va la soluci贸n del ejercicio 1
      
      ggplot(mtcars, aes(x=mpg,y=cyl)) +
        geom_point(color = "green")
      '''

Resultado:

```{r ejercicio2, exercise=TRUE}
# Cambia al color de los puntos a verde
ggplot(mtcars, aes(x=mpg,y=cyl)) +
  geom_point()
```

```{r ejercicio2-solution}
ggplot(mtcars, aes(x=mpg,y=cyl)) +
  geom_point(color= "green")
```

Al proporcionar hints y tambi茅n la soluci贸n en un ejercicio, 茅sta aparecer谩 dentro del apartado de las hints y se mostrar谩 al final.

C贸digo:

    Este es el c贸digo del chunk del ejercicio:

      '''{r ejercicio3, exercise=TRUE}
      #Aqu铆 va el ejercicio:
      
      # Cambia al color de los puntos a verde
      ggplot(mtcars, aes(x=mpg,y=cyl)) +
        geom_point()
      '''

    Este es el c贸digo del chunk del hint:

      '''{r ejercicio3-hint}
      #Aqu铆 va el hint
      
      # Usa el par谩metro color dentro de geom_point()
      '''
    Este es el c贸digo del chunk de la soluci贸n:

      '''{r ejercicio3-solution}
      #Aqu铆 va la soluci贸n del ejercicio 1
      
      ggplot(mtcars, aes(x=mpg,y=cyl)) +
        geom_point(color = "green")
      '''

Resultado:

```{r ejercicio3, exercise=TRUE, highlight=TRUE}
# Cambia al color de los puntos a verde
ggplot(mtcars, aes(x=mpg,y=cyl)) +
  geom_point()
```

```{r ejercicio3-hint}
# Usa el par谩metro color dentro de geom_point()
```

```{r ejercicio3-solution}
ggplot(mtcars, aes(x=mpg,y=cyl)) +
  geom_point(color= "green")
```

### 驴C贸mo puedo ayudar al usuario? *(completar c贸digo y aviso de fallas)*

Por defecto, en cada ejercicio de c贸digo se tiene que`exercise.completion = TRUE` lo cual hace que mientras el usuario escribe una funci贸n de R (por ejemplo `mean)`, se muestren los par谩metros que requiere esa funci贸n (por ejemplo `x = , na.rm =` ).

Ejemplo: Cuando `exercise.completion = FALSE`.

```{r ejercicio5, exercise = TRUE, exercise.completion = FALSE, exercise.cap = "Sin ayuda para completar"}
#Intenta volver a llenar los argumentos de la funci贸n seq para NO ver en marcha
# exercise.completion (pues pusimos exercise.completion = FALSE)
x <- seq(0,10,2)
```

Por otro lado `exercise.diagnostics = TRUE` sirve para avisar al usuario cuando cometa un error en el c贸digo, por ejemplo `function(` es un error de codificaci贸n porque falta un par茅ntesis, si se escribe `exercise.diagnostics = FALSE` no se le dar谩 aviso al usuario que es un error.

-   `exercise.diagnostics = TRUE`

    ![](images/learnr3.png)

-   `exercise.diagnostics = FALSE`

    ![](images/learnr4.png)

Ejemplo: Cuando `exercise.diagnostics = FALSE`.

```{r ejercicio6, exercise = TRUE, exercise.diagnostics = FALSE, exercise.cap = "Sin diagn贸stico"}
# Observa que no aparece del lado izquerdo alguna advertencia de los siguientes errores:
function(
  
mean(
```

### Examen sorpresa え

```{r cuestionario9}
quiz(
  question("驴Qu茅 opci贸n del chunk se usa para que aparezca el bot贸n de `Start over` en cada ejercicio?",
           answer("`exercise.diagnostics`",message="Esta opci贸n es para advertirte sobre los errores de c贸digo en el editor."),
           answer("`exercise.start-over`",correct=TRUE),
           answer("`exercise.cap`",message="Esta opci贸n te permite cambiar el t铆tulo del ejercicio."),
           allow_retry=TRUE,
           correct = "隆S铆 que has puesto atenci贸n!",
           try_again = "Vuelve a leer la secci贸n anterior",
           submit_button = "Yo s茅 que s茅",
           try_again_button = "Int茅ntalo de nuevo "),
  question("La opci贸n `exercise.timelimit` sirve para limitar el tiempo que tiene el usuario para responder.",
           answer("Verdadero"),
           answer("Falso", correct=TRUE),
           allow_retry=TRUE,
           correct = "隆S铆 que has puesto atenci贸n!",
           try_again = "![](https://images.gr-assets.com/hostedimages/1512773227ra/24632217.gif)",
           submit_button = "Yo s茅 que s茅",
           try_again_button = "Int茅ntalo de nuevo")
)

```

## Cuestionarios

Si el tema de tutorial requiere reforzar conceptos o relaciones (en lugar de c贸mo hacer el c贸digo) puedes lograrlo incluyendo preguntas (o cuestionarios) de opci贸n m煤ltiple sobre el tema.

### Preguntas

Para crear una pregunta de opci贸n m煤ltiple usa la siguiente estructura:

**C贸digo**

```{r,echo=TRUE,eval=FALSE}
{r quiz}
quiz(
  question("Aqu铆 escribes la pregunta",
    allow_retry = TRUE, #Que lo vuelva a intentar
    answer("Opci贸n 1"),
    answer("Opci贸n 2"),
    answer("Opci贸n correcta", correct = TRUE)
  ),
  question("Pueses incluir mensajes para cada opci贸n y m谩s de una opci贸n correcta.",
    allow_retry = TRUE,
    answer("Una opci贸n correcta", correct = TRUE,message = "Mensaje adicional"),
    answer("Una opci贸n incorrecta", message = "Mensaje adicional"),
    answer("Otra opci贸n incorrecta"),
    answer("Otra opci贸n correcta", correct = TRUE)
  )
)
```

**Resultado**

```{r cuestionario6}
quiz(
  question(text = "Aqu铆 escribes la pregunta",
    allow_retry = TRUE,
    answer("Opci贸n 1"),
    answer("Opci贸n 2"),
    answer("Opci贸n correcta", correct = TRUE)
  ),
  question("Pueses incluir mensajes para cada opci贸n y m谩s de una opci贸n correcta.",
    allow_retry = TRUE,
    answer("Una opci贸n correcta", correct = TRUE,message = "Mensaje adicional"),
    answer("Una opci贸n incorrecta", message = "Mensaje adicional"),
    answer("Otra opci贸n incorrecta"),
    answer("Otra opci贸n correcta", correct = TRUE)),
  caption = "Ejemplo"
)
```

### Par谩metros de `question`

**Lo b谩sico:**

-   `text = "驴Esta es la pregunta?"` El par谩meto `text` contiene una cadena de caracteres o **strings** (es decir texto dentro de "comillas").

-   `answer("Esta es la respuesta correcta", correct = TRUE, message = "Aqu铆 le dices al usuario algo bonito por haber acertado ;)")`

    -   La funci贸n answer es un par谩metro de question y contiene a su vez los siguentes par谩metros:

        -   `text` par谩metro de caracteres que definen una posible respuesta a la pregunta

        -   `correct` un par谩metro boolean(o), es decir, TRUE o FALSE. Por defecto, `correct = FALSE` lo cual indica que esa posible respuesta NO ES correcta.

        -   `message` par谩metro de caracteres para dar un mensaje adicional al usuario sobre la opci贸n de respuesta que eligi贸.

-   **驴Quieres permitir que el usuario vuelva a intentar responder?**

    Puedes hacerlo con `allow_retry` que recibe un boleano. Por defecto (es decir, sino se especifica el valor del par谩metro) `allow_retry = FALSE` lo cual implica que NO habr谩 posibilidad de que el usuario intente contestar de nuevo (si es que contest贸 incorrectamente).

-   **驴Quieres personalizar el mensaje que aparece al contestar correcta o incorrectamente?**

    Usa `correct = "Este mensaje saldr谩 cuando aciertes en tu respuesta"` .

    `incorrect = "Este mensaje saldr谩 cuando NO aciertes en tu respuesta"`, **solo aparecer谩 si** `allow_retry = FALSE`

    `try_again = "Este mensaje saldr谩 cuando NO aciertes en tu respuesta y tengas la opci贸n de intentarlo de nuevo"`

    **驴C贸mo cambiar los botones?**

-   `submit_button = "Mensaje que tiene el bot贸n para confirmar tu respuesta"`, por defecto `submit_button = "Submit Answer"`

-   `try_again_button = "Mensaje que tiene el bot贸n para hacer otro intento"` (siempre y cuando te equivoques), por defecto`try_again_button = "Try Again"`

-   `random_answer_order = TRUE` es un par谩metro que permite cambiar el orden de las opciones de respuesta cada que vuleves a realizar el tutorial. **Es decir dar click en `Start over`**

    Ejemplo:

```{r cuestionario7, echo = TRUE}
  question(
    text = "驴Cu谩l par谩metro sirve para cambiar el mensaje del bot贸n para intentar de nuevo?",
    allow_retry=TRUE,
    answer("incorrect",message = "No sabes. Este sirve para dar un mensaje cuando te equivocas"),
    answer("submit_button",message = "No sabes. Aqu铆 cambias de mesnaje de confirmaci贸n"),
    answer("try_again_button", correct = TRUE, message = "隆Maravillos@ estudiante!"),
    correct = "隆S铆 que has puesto atenci贸n!",
    try_again = "Vuelve a leer la secci贸n de par谩metros",
    submit_button = "Yo s茅 que s茅",
    try_again_button = "Acu茅rdate de m铆",
    random_answer_order = TRUE
  )
```

Si queremos hacer m煤ltiples preguntas usamos la funci贸n `quiz` que tenga por par谩metros las diferentes `quiestion`s.

```{r cuestionario8, echo=TRUE}
quiz( 
  question(
    text = "驴Cu谩les son los par谩metros de answer?",
    answer("text",correct = TRUE),
    answer("correct",correct = TRUE),
    answer("message", correct = TRUE), 
    correct = "隆Muy bien!",
    submit_button = "Yo s茅 que s茅", 
    allow_retry = TRUE,
    try_again_button = "Intenta de nuevo",
    try_again = "A lo mejor no leiste bien la pregunta"
    ),
  question(
    text = "驴Qu茅 funci贸n se usa para armar un cuestionario?",
    allow_retry=TRUE,
    answer("quiz", correct = TRUE, message = "Vas por buen camino"),
    answer("questions",message = "Esta funci贸n no existe en learnr"),
    answer("quizes", message = "Esta funci贸n no existe en learnr"), 
    correct = "Sigue adelante",
    try_again = "Intenta de nuevo",
    submit_button = "Yo s茅 que s茅",
    try_again_button = "Date otra oportunidad"
  ),
caption = "驴Aprendiste?"
)
```

### Para incluir s铆mbolos matem谩ticos en las preguntas:

Para poder hacer uso de f贸rmulas dentro de las preguntas, solo ponemos las f贸rmulas entre signos de pesos (\$) en sint谩xis de $\mathbb{L_aT_ex}$ y dos backslash `\\`, por ejempo:

```{r quiz4, echo=TRUE}
question(
  
  text = "驴A qu茅 distribuci贸n corresponde la siguiente funci贸n de densidad? $f(x) = \\frac{e^{\\frac{-(x-\\mu)^2}{2\\sigma^2}}}{\\sqrt{2\\pi\\sigma^2}}$",
  
  answer("Gaussiana", correct = TRUE),
  answer("Normal", correct = TRUE), 
  
  
#Este par谩metro permite al usuario que aun cuando seleccione una respuesta correcta de muchas, 茅sta se evaluar谩 como correcta (a diferencia de "checkbox" en la que hay que marcar todas las respuestas correctas para que se eval煤e como correcta).
  type = "radio", 
  correct = "隆驴Eres Gauss!?",
  submit_button = "Conf铆a en t铆"
)
```

## Incluir elementos externos

**驴Qu茅 otros elementos se pueden incluir para hacer m谩s llamativo mi tutorial?**

### Im谩genes

Como en cualquier tipo de archivo .Rmd puedes incluir im谩genes. Sin embargo, en este caso debe crearse una carpeta dentro del directorio del tutorial con el nombre **images** y guardar en esta las im谩genes a incluir y referenciarlas a dicha carpeta. Por ejemplo:

`![Mi imagen](images/mi_imagen.png)`

Tambi茅n se pueden incluir im谩genes (y gifs) de internet:

`![Imagen](https://link-a-la-imagen)`

![](https://media.tenor.com/images/36f456475f98b4770b4e9ad1016934f1/tenor.gif){style="figure-align: center"}

### Videos

Se usa la misma sint谩xis que para incluir una imagen de internet pero ahora escribimos el link al video. Solo pueden incluirse links a videos de YouTube o vimeo.

Lo atractivo de esta caracter铆stica es que el video se reproducir谩 dentro del mismo archivo del tutorial (no te redirige a otra p谩gina).

### Links

Si deseas a帽adir material de consulta que se encuentra en internet, puedes incluir referencias dentro del tutorial que te lleven a dichas p谩ginas web.

`[Texto de la referencia](https://link-a-la-pagina-de-referencia)`

```{r cuestionario2}
question(text="驴Cu谩l de los siguientes comandos sirve para a帽adir un video dentro del tutorial?",
         answer("`![video](images/mi_imagen.png)`"),
         answer("`![video](https://link-al-video`)",correct=TRUE),
         answer("`https://link-al-video`"),
         allow_retry=TRUE,
         correct = "隆Muy bien!",
        try_again = "Te est谩s confundiendo",
        submit_button = "Confirmar",
        try_again_button = "Intentalo de nuevo")
```

## Dise帽o del tutorial

**驴C贸mo cambiar la apariencia de mi tutorial?**

### Temas

De manera simple, podemos hacer que el tutorial se vea con diferente estilo poniendo en el encabezado YAML `theme: "default"`, donde default se puede cambiar por alguno de los temas predeterminados, siguientes:

|              |              |               |
|:------------:|:------------:|:-------------:|
| `"default"`  | `"readable"` |   `"paper"`   |
| `"cerulean"` | `"spacelab"` | `"sandstone"` |
| `"journal"`  |  `"united"`  |  `"simplex"`  |
|  `"flatly"`  |  `"cosmo"`   |   `"yeti"`    |
|  `"darkly"`  |  `"lumen"`   |               |

```{r cuestionario3}
question(
  text = "驴Cu谩l es la forma correcta de cambiar el tema del tutorial a `flatly`?",
  allow_retry = TRUE,
  answer("![](images/learnr11.png)", correct = TRUE),
  answer("![](images/learnr12.png)"),
  answer("![](images/learnr13.png)"),
    correct = "隆Perfecto!",
    try_again = "Intenta de nuevo",
    submit_button = "Lo lograr茅",
    try_again_button = "Quiero salir bien",
    random_answer_order = TRUE
)
```

### Avance progresivo

Es posible que alg煤n tema del tutorial sea muy extenso y revelar todo el contenido a la vez no sea lo ideal. **驴C贸mo podemos ir mostrando el contenido de un tema por partes?**

Para mostrar los subt贸picos del tutorial (aquellos que tienen \#\#\#) uno a la vez dando click al bot贸n **Continue** para avanzar, tenemos dos opciones:

-   Aplicarlo a todo el documento activando la opci贸n `progressive: true` en el encabezado YAML.

        ---
        title: "Hello, Tutorial!"
        output: 
          learnr::tutorial:
            progressive: true
        runtime: shiny_prerendered
        --- 

-   Aplicarlo solo a algunos subt贸picos agregando `{data-progressive=TRUE}` adelante de cada encabezado de tema (con \#\# ) que queremos revelar de manera progresiva.

        ## Tema {data-progressive=TRUE}

```{r cuestionario4}
question(
  text = "`data-progressive=TRUE` sirve para revelar el contenido de los subtitulos poco a poco...",
  allow_retry = TRUE,
  answer("En todo el documento"),
  answer("En alg煤n tema del documento", correct = TRUE),
  answer("En todos los temas subsecuentes a esta especificaci贸n"),
    correct = "隆Perfecto!",
    try_again = "Intenta de nuevo",
    submit_button = "Lo lograr茅",
    try_again_button = "Quiero salir bien"
)
```

Una vez activada la opci贸n `progressive` no podremos avanzar en el tutorial sin contestar los ejercicios de cada secci贸n. Para permitir el avance sin contestar los ejercicios, podemos agregar `allow_skip: true` en el encabezado YAML.

    ---
    title: "Hello, Tutorial!"
    output: 
      learnr::tutorial:
        progressive: true
        allow_skip: true
    runtime: shiny_prerendered
    --- 

Para solo permitir avanzar en algunas secciones y no en otras sin contestar, podemos especificar`{data-allow-skip=TRUE}` o `{data-allow-skip=FALSE}` en cada tema.

```{r cuestionario5}
question(
  text = "驴Con qu茅 comando se permite al usuario avanzar en el tutorial sin haber contestado los ejercicios? (en todo el documento)",
  allow_retry = TRUE,
  answer("data-allow-skip=TRUE"),
  answer("allow_skip: false"),
  answer("allow_skip: true", correct = TRUE),
    correct = "隆Perfecto!",
    try_again = "Intenta de nuevo",
    submit_button = "Lo lograr茅",
    try_again_button = "Quiero salir bien"
)
```

### Otro formato del tutorial

**驴Es posible incluir ejercicios en otros tipos de documentos que no sean un tutorial (`learnr::tutorial`)?**

隆S铆! Por ejemplo, durante la presentaci贸n de un tema en clase, usando una presentaci贸n tipo [`slidy`](https://garrettgman.github.io/rmarkdown/slidy_presentation_format.html). Para esto, se debe incluir `runtime: shiny_prerendered` en el encabezado YAML del documento de la presentaci贸n:

    # title: "Habits"
    # author: John Doe
    # date: March 22, 2005
    # output: slidy_presentation
    # runtime: shiny_prerendered

Antes de incluir un chunk de ejercicio en la presentaci贸n solo cargamos la paqueter铆a `learnr`.

## 驴Una 煤ltima pregunta?

```{r cuestionario10}
question("驴Usar谩s `learnr` despu茅s de revisar este tutorial?",
         type="radio",
         answer("隆Por supuesto!", correct=TRUE),
         answer("隆Claro que s铆!", correct=TRUE),
         correct="![](https://media1.tenor.com/images/4869c336eb8589454f7dc3582d623eb5/tenor.gif?itemid=6247482)",
         submit_button="Ya es lo 煤ltimo")

```

## Algo extra: Calificar c贸digo

**驴Es posible dar retroalimentaci贸n m谩s personal en los ejercicios de c贸digo?**

隆S铆! Si deseas que los usuarios tengan retroalimentaci贸n cuando cometan alg煤n error en los ejercicios de c贸digo (en t茅rminos del objetivo de aprendizaje en el ejercicio), puedes implementarlo con lo que veremos en esta secci贸n.

### Instalar `gradethis`

La paqueter铆a que usaremos para este fin se llama `gradethis` y te sugerimos hacer lo siguiente para tenerla instalada en tu computadora:

1.  Si ya instalaste `learnr` con el comando `install.packages("learnr")` entonces utiliza `remove.packages("learnr")` para desinstalarlo.

    Sea afirmativo o no lo anterior, haz lo siguiente:

2.  Usa el comando `remotes::install_github("rstudio/learnr")` para tener la versi贸n en desarrollo de `learnr`.

3.  Usa el comando `remotes::install_github("rstudio-education/gradethis", force = TRUE)` para instalar `gradethis`.

4.  Cargar la paqueter铆a con el comando `library(gradethis)` despu茅s de cargar `learnr`, ambas dentro del chunk inicial (que se puede llamar setup).

### Esquema

![](images/gradethis.png){width="692"}

Existen dos formas principales para evaluar el ejercicio de c贸digo:

1.  Evaluar el c贸digo.

    En la paqueter铆a `gradethis` esto se puede hacer con la funci贸n `grade_this_code()`. Esta funci贸n reemplaza a la funci贸n `grade_code()` ( que se menciona [en la documentaci贸n de](https://rstudio.github.io/learnr/exercises.html#Exercise_Checking)`learnr`).

2.  Evaluar el resultado que se deriva del c贸digo.

    Con la funci贸n `grade_this()` se puede personalizar el criterio de evaluaci贸n con respecto al resultado. **Por ejemplo**: si el resultado es mayor al esperado, se puede implementar unmensaje que diga "el resultado correcto es menor al obtenido" .

    Esta funci贸n, como en el punto anterior, reemplaza a `grade_result()`.

### grade_this_code()

**Ejemplo:**

Implementemos un ejercicio donde el usuario tenga que sumar 2+2 (= ) en el que se evaluar谩 que el c贸digo sea `2+2`.

1.  Crear el chunk de ejercicio.

2.  Crear el chunk solution.

    ```{r gejercicio-solution, exercise.reveal_solution = FALSE}
    2+2
    ```

3.  Crear el chunk check:

    El chunk de ejercicio se llama `gejericicio`, por lo que el nombre del chunk check debe ser `gejercicio-check`.

    Dentro del chunk siguiente se usar谩 la funci贸n `grade_this_code()`.

        '''{r gejercicio-check}

        grade_this_code()

        '''

    ```{r gejercicio-check}
    gradethis::grade_this_code()
    ```

Resultado:

```{r gejercicio, exercise = TRUE, exercise.cap = "Suma"}
# Suma dos m谩s dos

```

-   Nota que no importa si pones espacios en el c贸digo o pones el c贸digo en otras l铆neas (diferentes a las del chunk solution).

-   Nota que se a帽ade un bot贸n que se llama `Submit answer` para que el c贸digo sea evaluado.

-   Intenta escribir `3+1` , `2*2`, `4*1,` ... y nota que esa no es la respuesta correcta aun cuando el resultado sea el mismo (4).

### Par谩metros

Se pueden modificar los mensajes que aparecen cuando se comete un error o un acierto con los par谩metros de `grade_this_code()`:

-   `correct` recibe una cadena de caracteres la cual aparece cuando el c贸digo del usuario y el c贸digo de la soluci贸n coinciden.

    Por ejemplo `correct = "Haz hecho la suma que deseo"`

-   `incorrect` recibe una cadena de caracteres la cual aparece cuando el c贸digo del usuario y el c贸digo de la soluci贸n NO coinciden.

    Por ejemplo `incorrect = "Intenta sumar dos n煤meros iguales y que el resultado sea 4"`

-   `allow_partial_matching` puede ser `TRUE` o `FALSE`. Por defecto es `FALSE` y cuando `allow_partial_matching = TRUE` entonces se permite que aparezcan mensajes que indican d贸nde no coincide el c贸digo con la soluci贸n.

**Ejemplo** del ejemplo anterior tenemos:

    '''{r gejercicio-check}

    grade_this_code(correct = "Muy bien", 
                    incorrect = "Intenta sumar dos n煤meros iguales y que el resultado sea 4", 
                    allow_partial_matching = TRUE)

    '''

```{r gejercicio1, exercise = TRUE, exercise.cap = "Suma"}
# Suma dos m谩s dos

```

```{r gejercicio1-solution, exercise.revela_solution=FALSE}
2+2
```

```{r gejercicio1-check}
grade_this_code(correct = "Muy bien", 
                incorrect = "Intenta sumar dos n煤meros iguales y que el resultado sea 4", 
                allow_partial_matching = TRUE)
```

### 驴Te qued贸 duda?

```{r gcuestionario}
quiz(
  question("Si el chunk del ejercicio se llama 'ejercicio', 驴En qu茅 chunk debe ir la funci贸n `grade_this_code()`?",
    allow_retry = TRUE,
    answer("ejercicio-solution", message = "No, aqu铆 debe ir el c贸digo esperado"),
    answer("ejercicio-check", correct = TRUE),
    answer("ejercicio-cechk", message = "Revisa la escritura"),
    correct = "隆S铆 que has puesto atenci贸n!",
    try_again = "Vuelve a leer la secci贸n anterior",
    submit_button = "Yo s茅 que s茅",
    try_again_button = "Int茅ntalo de nuevo "
  ),
  
  question("驴Qu茅 par谩metro permite al usuario saber en qu茅 parte del c贸digo se pudo haber equivocado?", 
    answer("allow_retry"),
    answer("allow_partial_matching", correct = TRUE),
    answer("incorrect"),
    allow_retry = TRUE,
    correct = "隆S铆 que has puesto atenci贸n!",
    try_again = "Vuelve a leer la secci贸n anterior",
    submit_button = "Yo s茅 que s茅",
    try_again_button = "Int茅ntalo de nuevo "
  ),
  
  question("驴Para que sirve el par谩metro correct?",
    answer("Para permitir que el usuario tenga correcto lo que haga siempre"),
    answer("Para dar a conocer al usuario un mensjae cuando el su c贸digo coincide con el de la soluci贸n", correct = TRUE),
    answer("Para estar felices"),
    allow_retry = TRUE,
    correct = "隆S铆 que has puesto atenci贸n!",
    try_again = "Vuelve a leer la secci贸n anterior",
    submit_button = "Yo s茅 que s茅",
    try_again_button = "Int茅ntalo de nuevo "
  ), 
  caption = "Pongamos a prueba lo que haz aprendido"
)
```

### grade_this()

**Ejemplo:**

Implementemos un ejercicio donde el usuario tenga que obtener el resultado de sumar 2+2 (= ) y se evaluar谩 que dicho resultado sea `4`.

1.  Crear el chunk de ejercicio.

2.  Crear el chunk solution.

    ```{r gtejercicio-solution, exercise.reveal_solution = FALSE}
    2+2
    ```

3.  Crear el chunk check:

    El chunk de ejercicio se llama `gtejericicio`, por lo que el nombre del chunk check debe ser `gtejercicio-check`.

    En el chunk siguiente se usar谩 la funci贸n `grade_this()`. Dentro de esta funci贸n se debe definir la l贸gica que se usa para evaluar el ejercicio. En el ejemplo se muestra una forma de comparar el resultado del usuario (`.result`) con el resultado correcto (`.solution`) y si son iguales se marca como correcto el ejercicio.

    Esta es solo una forma de varias que se pueden implementar para evaluar el ejercicio, las cuales veremos m谩s a detalle en la secci贸n siguiente.

        '''{r gtejercicio-check}

        grade_this(
          if(.result == .solution){
            pass("Muy bien")
          } else{
            fail("El resultado debe ser igual a 2+2")
          }
        )

        '''

    ```{r gtejercicio-check}
    gradethis::grade_this(
      if(.result == .solution){
        pass("Muy bien")
      } else{
        fail("El resultado debe ser igual a 2+2")
      }
    )
    ```

Resultado:

```{r gtejercicio, exercise = TRUE, exercise.cap = "Suma"}
# Obten el resultado de sumar dos m谩s dos

```

-   Notemos que no importa si el resultado correcto (`4`) se obtiene de otra forma diferente (`3+1`, `4*1`, `4+0`) a la propuesta. Mientras sea `4` el ejercicio se marca como correcto.

### Funciones para definir la l贸gica de evaluaci贸n

Dentro de `gradethis` existen funciones que se puede usar en `grade_this()` para ayudar a definir la forma de evaluar los ejercicios de c贸digo.

**Funciones necesarias**

+-------------+----------------------------------------------------------------------------------------------------+
| Funci贸n     | Descripci贸n                                                                                        |
+=============+====================================================================================================+
| `pass()`    | En esta funci贸n se escribe el mensaje que aparecer谩 cuando se marque como correcto el ejercicio.   |
+-------------+----------------------------------------------------------------------------------------------------+
| `fail()`    | En esta funci贸n se escribe el mensaje que aparecer谩 cuando se marque como incorrecto el ejercicio. |
+-------------+----------------------------------------------------------------------------------------------------+
| `.result`   | Hace referencia al resultado que se obtiene con el c贸digo del usuario.                             |
+-------------+----------------------------------------------------------------------------------------------------+
| `.solution` | Hace referencia al resultado correcto que se obtiene con el c贸digo del chunk `solution`.           |
+-------------+----------------------------------------------------------------------------------------------------+

**Funciones auxiliares para evaluar**

+-------------------+---------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------+
| Funci贸n           | Descripci贸n                                                                                                                                                                                                         |
+===================+=====================================================================================================================================================================================================================+
| `pass_if()`       | En el par谩metro `cond` se recibe la condici贸n que debe cumplir .`result` para que el ejercicio sea correcto y en el par谩metro `message` se escribe el mensaje que aparecer谩 cuando la condici贸n anterior se cumpla. |
+-------------------+---------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------+
| `fail_if()`       | En el par谩metro `cond` se recibe la condici贸n con la cual se evaluar谩 como incorrecto el ejercicio y el par谩metro `message` recibe el mensaje que aparecer谩 cuando dicha condici贸n se cumpla.                       |
+-------------------+---------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------+
| `pass_if_equal()` | Compara los par谩metros `y` y `x` y evalua el ejercicio como correcto cuando estos son iguales, devolviendo el mensaje que se le pase con el par谩metro `message`. Por default `y` es `.solution` y `x` es `.result`. |
+-------------------+---------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------+
| `fail_if_equal()` | Compara los par谩metros `y` y `x` y evalua el ejercicio como incorrecto cuando estos son iguales, devolviendo el mensaje que se le pase en el par谩metro `message`. Por default `y` es `.solution` y x es `.result`.  |
+-------------------+---------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------+

**Ejemplo**:

En el ejercicio `gt_ejercicio` se pide obtener cu谩ntas observaciones de la tabla `penguins` tienen el valor `Torgensen` en el atributo `island`.

    '''{r gt_ejercicio-check}
    grade_this({
        # Se evalua autom谩ticamente como correcto si .result es igual a .solution
        pass_if_equal(message="隆Muy bien!")
        
        # Se evalua como incorrecto cuando el n煤mero que se especific贸 no es el         #  correcto y se muestra un mensaje personalizado.
        fail_if_equal(sum(penguins$island=="Biscoe"), message = "No son las observaciones de la isla correcta.")
        fail_if_equal(sum(penguins$island=="Dream"), message = "No son las observaciones de la isla correcta.")
        
        # Se muestra el mensaje autom谩tico de retroalimentaci贸n cuando existe           #  cualquier otro error en el c贸digo.
        fail()
      })
    '''

Resultado:

```{r gt_ejercicio, exercise=TRUE, exercise.cap = "pass_if_, fail_if"}
# Usando los datos de la tabla penguins, 驴cu谩ntos ping眉inos provienen de la isla Torgersen?

```

```{r gt_ejercicio-solution}
sum(penguins$island=="Torgersen")
```

```{r gt_ejercicio-check}
grade_this({
    # Se evalua autom谩ticamente como correcto si .result es igual a .solution
    pass_if_equal(message="隆Muy bien!")
    
    # Se evalua como incorrecto cuando el n煤mero que se especific贸 no es el         #  correcto y se muestra un mensaje personalizado.
    fail_if_equal(sum(penguins$island=="Biscoe"), message = "No son las
                  observaciones de la isla correcta.")
    fail_if_equal(sum(penguins$island=="Dream"), message = "No son las 
                  observaciones de la isla correcta.")
    
    # Se muestra el mensaje autom谩tico de retroalimentaci贸n cuando existe           #  cualquier otro error en el c贸digo.
    fail()
  })
```

-   Intenta escribiendo `sum(penguins$island=="Biscoe")` o `168` para ver el mensaje.

Tambi茅n se puede definir una l贸gica para evaluar usando cla煤sulas `if()` y `else()` como en el primer ejemplo de esta secci贸n y usar `pass()` y `fail()` para indicar cu谩ndo se evalua como correcto y cu谩ndo como incorrecto.

![](https://www.youtube.com/watch?v=fIqmew1pF6I)

**Ejemplo**:

*C贸digo del chunk-check*

```{r echo=TRUE,eval=FALSE, message=FALSE,warning=FALSE}
grade_this({
  # 驴El resultado es un objeto data.frame?
  if (!inherits(.result, "data.frame")) {
    fail("Tu resultado debe ser un data.frame")
  }

  # Ahora que sabemos que .result es un data.frame veremos si tiene la especie correcta...
  if (.result$species[1] != "Gentoo") {
    fail("Las observaciones deben ser de la especie Gentoo.")
  }

  # ...ahora comprobamos que la isla sea la correcta
  if (.result$island != "Biscoe") {
    fail("Las observaciones deben ser de la isla Biscoe.")
  }
  
  # ...comprobamos que el n煤mero de observaciones es el que debe ser
  if (nrow(.result) != 124) {
    fail("La tabla debe tener 124 observaciones.")
  }
  # ahora que revisamos las condiciones anteriores mostramos el mensaje de que el resultado es correcto
  pass(message="Ahora ya sabes filtrar datos de una tabla.")
})
```

```{r ejercicio_if-check, echo=TRUE,message=FALSE,warning=FALSE}
grade_this({
  # 驴El resultado es un objeto data.frame?
  if (!inherits(.result, "data.frame")) {
    fail("Tu resultado debe ser un data.frame")
  }

  # Ahora que sabemos que .result es un data.frame veremos si tiene la especie correcta...
  if (.result$species[1] != "Gentoo") {
    fail("Las observaciones deben ser de la especie Gentoo.")
  }

  # ...ahora comprobamos que la isla sea la correcta
  if (.result$island != "Biscoe") {
    fail("Las observaciones deben ser de la isla Biscoe.")
  }
  
  # ...comprobamos que el n煤mero de observaciones es el que debe ser
  if (nrow(.result) != 124) {
    fail("La tabla debe tener 124 observaciones.")
  }
  # all of the above checks have passed now.
  pass(message="Ahora ya sabes filtrar datos de una tabla.")
})

```

*Resultado*

```{r ejercicio_if, exercise=TRUE, exercise.cap="if - else", exercise.lines=5, message=FALSE,warning=FALSE}
# Crea una tabla que contenga solo los ping眉inos de la especie Gentoo 
#  y la isla Biscoe usando el data.frame penguins.
library(dplyr)
penguins %>% filter(species=="  ", island=="  ")

```

Intenta sustituir la respuesta anterior por:

-   `penguins %>% filter(species=="Adelie", island=="Dream")`

-   `head(penguins %>% filter(species=="Gentoo", island=="Biscoe"))`

-   `matrix(NA,3,2)`

## Un paso m谩s all谩: Un dise帽o personalizado

### CSS

Si quieres un dise帽o distinto (y propio) a los temas predeterminados puedes crear un archivo o secci贸n CSS para el tutorial.

Para ello considera el siguiente esquema que ayudar谩 a entender la sint谩xis que veremos m谩s adelante:

![](images/learnr14.png){width="694"}

CSS significa Cascading Style Sheets y es un lenguaje de hojas de estilo en cascada que separa los elementos del tutorial para que puedas modificarlos.

Los comandos principales para hacer estos cambios son:

-   `color` color de la letra

-   `font-family` tipo de letra

-   `font-size` tama帽o de la letra

-   `background-color` color de fondo

Un ejemplo para cambiar el body (ver imagen) es el siguiente:

![](images/learnr15.png)

Nota que el color lo pusimos en formato [hexadecimal](https://htmlcolorcodes.com/es/).

**Ejercicio.** Completa el siguiente c贸digo para cambiar el tipo de letra del nombre de los autores (Arial, Verdana, Helvetica [o la que se pueda](https://www.w3schools.com/css/css_font.asp)).

```{r ej2, exercise = TRUE, exercise.cap = "Mi primer CSS"}
__ {
  color: #8A6C0C;
  font-family: ______;
  font-size: 18px;
}

#Como esta no es sint谩xis de R, habr谩 un error al dar click en "Run Code"
```

Estas especificaones se pueden realizar en un archivo CSS o al inicio del tutorial.

Si lo pones al inicio del tutorial (no en un chunk de c贸digo) necesitar谩s un esquema como este:

    <style>

        h2 {
          color: #6D4101;
          font-size: 35px;
        }
        body {
          font-family: Verdana, Helvetica, sans-serif;
          font-size: 16;
          background-color: #EEF5F9;
          color: #000000;
        }
        
    </style>

Que se pone despu茅s del encabezado YAML.
